# Problem 1: Working with lists (6 points)

In class, we saw how to use *lists* to hold multiple data values. Lists are a powerful tool for handling data, and learning how to work with lists is one of the most important skills you'll learn in this class. Later in the course we will use our list skills to build complex data structures to work with genomic datasets. In this homework, we'll keep it simple and focus on some list basics.

### Creating a new list

A list in Python is just what it sounds like -- an *ordered* set of items. (There are unordered data types in Python, which we'll discuss later.) Lists can contain zero or more items, which are enclosed by square brackets and separated by commas:
```python
    []                # Empty list
    [567]             # List with one item
    [4,2,89,3]        # List of integers
    [4,2,"DNA", 3.14] # List of mixed data types - usually we don't mix data, but you can
```

To actually use a list, you need to assign it to a variable:

```python
    my_list = []  # create a blank list
```    

In the cell below, create a list of the numbers 1 through 10 called ```numbers```. (Be sure to run the cell by hitting shift-enter.)

<!-- grade_id: cell-b6f0aa255e523ffe, points: 0, grade: false, solution: true -->
```python
# YOUR ANSWER HERE

print(numbers)
```

<!-- grade_id: cell-5e35162fd9394c73, points: 1, grade: true, solution: false -->
```python
# Run this cell to check your answer. If it gives an error, your answer is not right.

assert numbers[-2] == 9
```

### Accessing items in a list: Indexing
To do something with an entire list, you just use the variable name of that list. But often you want to access individual items in the list. To do this, we use *indexing*. That is, you specify the position in the list you want to access, *starting with 0*. (Remember, computer scientists start counting at 0. This will be a common theme in your computational work.) Because lists are *ordered*, you access individual items by their position in the list.

So for example ```numbers[0]``` will give you the first item in the list ```numbers```, which is 1.

Less intuitively, you can use indexing by counting backwards from the end of the list. This can be a little confusing, because, unlike forward indexing, *counting backwards does not start from 0*. Look carefully at the following:

```python
numbers[-1] # last item
numbers[-3] # third from the last item
numbers[-len(numbers)] # first item - can you figure out why? Note the function nested within and index.
```

**Square brackets:** Somewhat confusingly, square brackets are used not only to define lists, but also for list indexing. *Don't confuse these two uses of square brackets!* They are not the same. When the brackets are placed right next to the variable name, they are being used for indexing:

```python
numbers[3]    # Gives you the 4th item in the list named 'numbers'.
numbers = [3] # Defines a brand new list (or overwrites an existing one!) with only one item. 
```

To practice indexsing, define a new list below, called ```more_numbers``` with four elements, the numbers 10, 11, 12, and 13 (in that order). Then, using indexing, calculate the sum of the **4th item** from ```numbers``` and the **second-to-last item** in ```more_numbers```.  Assign the answer to a new variable, `answer1`.

<!-- grade_id: cell-7e34ac727425af61, points: 0, grade: false, solution: true -->
```python
# Create the list more_numbers, then add the requested list items.

# YOUR ANSWER HERE

print(answer1)
```

<!-- grade_id: cell-aa088db9478ac0c8, points: 2, grade: true, solution: false -->
```python
assert answer1 == 16
assert more_numbers[-3] == 11
```

### Accessing items in a list: Slices

To access several adjacent items in a list, you can use a *slice*. A slice uses square brackets like an index, and also starts counting from 0:

```python
numbers[3]    # Indexing gives you the 4th item in the list.
numbers[3:7]  # Slicing gives you the 4th, 5th, 6th, and 7th items in the list.
numbers[0:3]  # First through 3rd items in the list
```

The numbering system for slices (inclusive start position, exclusive end position, counting starts from zero) can be confusing. But it's important to master it, becuase this numbering system will come up again and again.

One easy way to think about this is to subtract the first number from the last number. For `numbers[3:7]`, 7 - 3 = 4, the number of items grabbed by the slice.

Let's practice. In the cell below, use slicing to grab the numbers 11, 12, and 13 from the list `more_numbers`. Assign the result to the variable `answer2`. (What type of variable will `answer2` be after this? You can use `type()` or `print()` to check your guess.)

<!-- grade_id: cell-3afd954741094bb9, points: 0, grade: false, solution: true -->
```python
# Grab the numbers 11, 12, and 13 from more-numbers, assign to answer2

# YOUR ANSWER HERE

print(answer2)
```

<!-- grade_id: cell-2d9dc0dd28c1efa6, points: 2, grade: true, solution: false -->
```python
assert max(answer2) == 13
assert min(answer2) == 11
assert type(answer2) == list
```

### Combining list manipulations

In this last problem, you'll practice combining indexing with the `.append()` function to add a number from the list `more_numbers` to our list `numbers`.

Recall how in class we used `.append()` to add a **single** item to the end of a list. (`.pop()` removes items from the end of a list - though as you'll see later in the class, taking items off a list turns out to not be as useful as appending items.)

(What about putting items in the middle or beginning of a list? There are functions to do this, but using them is generally a bad idea. Unlike `.append()` and `.pop()`, inserting items in the middle of a list causes the position of every item in to change. Remember, lists are ordered, and store data by position. It is computationally slow to update the position of every item in lists with thousands of, say, gene expression measurements. Inserting things in the middle of lists makes bug-free coding harder.)

Example of using `.append()`

```python
nucleosides = ['deoxyadenosine', 'deoxyguanosine', 'deoxythymidine'] # We're missing one!
nucleosides.append('deoxycytidine') # Adds to end of list
```

In this problem, you'll combine list indexing with the `.append()` function. In the cell below, *use just one line of code* to take the number 11 from the list `more_numbers` (with indexing), and append it to the list `numbers`.

The substance of this problem is trivial, but implementing it in code is a little tricky: you have to take one operation (indexing) and nest it within a function. It's important to practice combinations like this. They are the basis of the powerful data manipulations we'll do using Python lists.

<!-- grade_id: cell-cb288e39c84b0ccc, points: 0, grade: false, solution: true -->
```python
# Append the number 11 (from the list more_numbers) to the list numbers
# Use just one line of code to do this

# YOUR ANSWER HERE

print(numbers)
```

<!-- grade_id: cell-7fc3fde8c5e33f6e, points: 2, grade: true, solution: false -->
```python
assert len(numbers) == max(numbers)
```

### Bonus material - 2 extra points

Here's something to try in the cell below. If you wanted to attach last 3 items of the list `more_numbers` to `numbers`, how would you do it? Try out your ideas in the cell below.

What happens if you try to combine a slice of `more_numbers` with `.append()`? It's not what you expect. (Hint: check the output of `print()`. Pay particular attention to the positions of the square brackets.)

Here are two hints for one way to do this: 

1. How did we combine two strings? Will this work for lists?
2. A variable can be defined in terms of itself plus another variable (though think carefully about whether you want to alter a variable, or just alter a copy of that variable):

```python
total = 10 # 10 dollars in your account
total = total - 1 # Subtract a dollar. Total is used to define a new value of total
```

Using the concepts of these two hints, you can write one line of code to append *some* of `more_numbers` to `numbers`.

<!-- grade_id: cell-d4236df66a6dff6e, points: 0, grade: false, solution: true -->
```python
# I've reset the list numbers to its original state. Using just one line of code,  
# try attaching the last three items from more_numbers, so that numbers is
# a list of integers from 1 to 13.

numbers = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]

# YOUR ANSWER HERE
print(numbers)
```

<!-- grade_id: cell-7e1c10f25add86a8, points: 2, grade: true, solution: false -->
```python
assert len(numbers) == 13
assert sum(numbers) == 91
```

For more on lists, here are some hepful links:

http://www.python-course.eu/python3_sequential_data_types.php

http://www.python-course.eu/python3_list_manipulation.php

https://docs.python.org/3/tutorial/introduction.html#lists