# Problem 1: Measures of Scale (10 points)

In this problem, we will write code to calculate the variance, standard deviation, and median absolute deviation.

First, load up the data and copy your `mean` and `median` functions from the previous problem below:

<!-- grade_id: cell-721d8dc00a482ef7, points: 0, grade: false, solution: true -->
```python
toy_data = [1, 1, 2, 4, 32]

exon_lengths = []
file = open('gene_table.txt')
header = file.readline()
for line in file:
    values = line.strip('\n').split('\t')
    exon_lengths.append(int(values[3]))
file.close()

# now paste in below your median and mean functions
# YOUR ANSWER HERE

# if you couldn't get mean and median working in the previous problem, uncomment the below lines to get
# pre-canned versions:
# from numpy import mean, median
```

## Variance and Standard Deviation (5 points)

Write a function to calculate the variance. Remember that the variance is simply the mean squared deviation, where the "deviation" is the difference between each value and the mean.

Hint: In Python, you can raise a number to a given power with `**`. (If one `*` means “multiplication”, two must mean “extra multiplication”, I guess :) So, you can write `4**2` to mean $4^2$, `7.5**-4` for $7.5^{-4}$, and so on. Don't forget operator precedence: `1-4**2` is -15, not 9. Use parentheses as needed.

<!-- grade_id: cell-fab32404573f12e9, points: 0, grade: false, solution: true -->
```python
def variance(values):
    mean_value = mean(values) # only calculate this once! Don't call mean(values) each time through the loop!!
    squared_deviations = []
    # fill in the squared_deviations list and then return its mean value
    # YOUR ANSWER HERE

print(variance(toy_data))
print(variance(exon_lengths))
```

<!-- grade_id: cell-261a66f099459080, points: 4, grade: true, solution: false -->
```python
assert 145 < variance(toy_data) < 145.5
assert 46.5 < variance(exon_lengths) < 47
```

Now write a function for standard deviation, which uses the `variance` function and then returns the square root of that. (Hint: remember that $\sqrt{x} = x^{0.5}$)

<!-- grade_id: cell-286941d6733edcda, points: 0, grade: false, solution: true -->
```python
def stdev(values):
    # YOUR ANSWER HERE

print(stdev(toy_data))
print(stdev(exon_lengths))
```

<!-- grade_id: cell-65914a61f3bbe1d8, points: 1, grade: true, solution: false -->
```python
assert 12 < stdev(toy_data) < 12.05
assert 6.8 < stdev(exon_lengths) < 6.9
```

Now write a function to calculate the median absolute deviation (where "deviation" is just the numerical difference between each data point and the median, and "absolute" means "take the absolute value of the deviation").

<!-- grade_id: cell-b05c81f43e8af13c, points: 0, grade: false, solution: true -->
```python
def mad(values):
    # YOUR ANSWER HERE

print(mad(toy_data))
print(mad(exon_lengths))
```

<!-- grade_id: cell-1065a83de48a415a, points: 5, grade: true, solution: false -->
```python
assert mad(toy_data) == 1
assert mad(exon_lengths) == 2
```