# Problem 0: Zipping lists and more complex conditional statements (6 points)

In problem 0 you'll practice using `zip()` and conditional statements, working again with our toy reporter gene data set.

**REMINDER**: You can create blank cells to write some scratch code if you need extra space. The autograde will ignore these. However, your actual answer code must go in the answer cells. 

## Problem 0.1: Warmup with `zip()` and `if`

First, a quick problem to practice the syntax:

You have some data: a list of transcript types, and a corresponding list of expression changes collected during an experiment. Your task is to identify transcripts that are 1) protein-coding and 2) whose expression decreased in the experiment.

To do this problem, you'll need to build an **`if`** statement that depends on two conditions: the gene is protein-coding **and** its expression went down. Sensibly enough, we test for the truth values of two conditions in Python using **`and`**:

```python
# Probably not the code for Google's self-driving car
color = 'red'
shape = 'octagon'

if color=='red' and shape=='octagon':
    print('STOP!')
else:
    print('GO!')
```

In the cell below, use `zip()` to simultaneously loop over two lists, `gene_types` and `expression`. Inside the block of the **`for`** loop, use an `if` statement with two conditions to identify protein-coding transcripts whose expression decreased (i.e., negative values). Append those expression values to the list `downreg_pc`. (For the other transcripts, do nothing.)

<!-- grade_id: cell-c7a567a6c2a36988, points: 0, grade: false, solution: true -->
```python
# I defined the lists for you
gene_types = ['protein_coding','antisense','protein_coding','protein_coding','miRNA','protein_coding','lincRNA']
expression = [-0.22, 0.77, 0.03, -0.62, 0.19, 0.23,-0.49]
downregulated_pc = [] # list of downregulated protein genes


# YOUR ANSWER HERE

# See your results:
print("Downregulated expression values:",downregulated_pc) 
```

<!-- grade_id: cell-8cbb4afc1e50045a, points: 2, grade: true, solution: false -->
```python
# Run to check your answer
assert len(downregulated_pc) == 2
assert downregulated_pc[0] == -0.22
assert downregulated_pc[1] == -0.62
```

## Problem 0.2: Reporter gene data version 3

In this problem, we return yet again to the red/green reporter gene experiment we saw in lecture 1 and problem set 1. The problem is the same one you were asked to complete in the homework, but this time you should solve it with `zip()`.

If you need to review the syntax, look back over the code you wrote for the Lecture 1 non-graded activity and homework problem 2.

Recall that your task in the reporter gene problem is to write a `for` loop that iterates over both `red` and `green` data at the same time, and divide the red value by its corresponding green value. The result should be appended to a new list called `norm_data`.

Instead of using the `range()` syntax, this time solve the problem using `zip()`. In the cell below, do the following:


1. Create an empty list called `norm_data` to hold the normalized values.

2. Write a `for` loop using `zip()` to iterate over both `red` and `green`. 

3. In the loop block (indented lines), divide the red value by the green value and `.append()` the result to the list `norm_data`.


<!-- grade_id: cell-a6f2c01674f7e7a4, points: 0, grade: false, solution: true -->
```python
red = [23, 145, 203, 235, 354, 456]
green = [5, 11, 6, 9, 8, 4]

# Write your code below. Use print() to display your results.
# YOUR ANSWER HERE
```

<!-- grade_id: cell-bbe8cb979757d0e2, points: 1, grade: true, solution: false -->
```python
assert len(norm_data) == len(red)
assert norm_data[2] < 33.834 and norm_data[0] == 4.6
```

## Problem 0.3: Reporter gene data version 4 - avoid dividing by 0

Zeros in our control data are another potential source of error in our reporter gene normalization code. If we got a blank measurement for one sample in the green channel, then our code above won't work – you can't divide red by green if green is zero.

So let's make our normalization code a little more sophisticated, by adding a conditional statement to check for zero values in the green channel.

We want a `for` loop that recognizes when green is zero and then executes an alternate instruction.

In the cell below, I've edited `green` to include a 0. Write a `for` loop to normalize red by green, as before. Your tasks is to use `if` and `else` inside the `for` loop block to do the following:

 *If* the value of green is greater than zero, normalize green by red and append the result to `normalized`.
 
 *Else*, print an error message: "Green value is zero." (And do nothing else - we skip that sample.)

<!-- grade_id: cell-f2ec1ae9626b7e33, points: 0, grade: false, solution: true -->
```python
red = [23, 145, 203, 235, 354, 456]
green = [5, 11, 6, 0, 8, 4]
norm_data = [] # clear the list - we'll reuse this variable name

# Write your code below.

# YOUR ANSWER HERE
print(norm_data)
# Visualize the result.
```

<!-- grade_id: cell-ed3ed698877fd4ce, points: 3, grade: true, solution: false -->
```python
assert len(norm_data) == 5
assert norm_data[2] < 33.834 and norm_data[-1] == 114.0
```